Colony demographics shape nest construction in Camponotus fellah ants

Author response:

The following is the authors’ response to the previous reviews.

Reviewer #1 (Public review):

This study investigates how ant group demographics influence nest structures and group behaviors of Camponotus fellah ants, a ground-dwelling carpenter ant species (found locally in Israel) that build subterranean nest structures. Using a quasi-2D cell filled with artificial sand, the authors perform two complementary sets of experiments to try to link group behavior and nest structure: first, the authors place a mated queen and several pupae into their cell and observe the structures that emerge both before and after the pupae eclose (i.e., "colony maturation" experiments); second, the authors create small groups (of 5,10, or 15 ants, each including a queen) within a narrow age range (i.e., "fixed demographic" experiments) to explore the dependence of age on construction. Some of the fixed demographic instantiations included a manually induced catastrophic collapse event; the authors then compared emergency repair behavior to natural nest creation. Finally, the authors introduce a modified logistic growth model to describe the time-dependent nest area. The modification introduced parameters that allow for age-dependent behavior, and the authors use their fixed demographic experiments to set these parameters, and then apply the model to interpret the behavior of the colony maturation experiments. The main results of this paper are that for natural nest construction, nest areas, and morphologies depend on the age demographics of ants in the experiments: younger ants create larger nests and angled tunnels, while older ants tend to dig less and build predominantly vertical tunnels; in contrast, emergency response seems to elicit digging in ants of all ages to repair the nest.

The experimental results are solid, providing new information and important insights into nest and colony growth in a social insect species. As presented, I still have some reservations about the model's contribution to a deeper understanding of the system. Additional context and explanation of the model, implications, and limitations would be helpful for readers.

We sincerely thank Reviewer #1 for the time and effort dedicated to our manuscript's detailed review and assessment. The new revision suggestions were constructive, and we have provided a point-by-point response to address them.

Reviewer #2 (Public review):

I enjoyed this paper and its examination of the relationship between overall density and age polyethism to reduce the computational complexity required to match nest size with population. I had some questions about the requirement that growth is infinite in such a solution, but these have been addressed by the authors in the responses and the updated manuscript. I also enjoyed the discussion of whether collective behaviour is an appropriate framework in systems in which agents (or individuals) differ in the behavioural rules they employ, according to age, location, or information state. This is especially important in a system like social insects, typically held as a classic example of individual-as-subservient to whole, and therefore most likely to employ universal rules of behaviour. The current paper demonstrates a potentially continuous age-related change in target behaviour (excavation), and suggests an elegant and minimal solution to the requirement for building according to need in ants, avoiding the invocation of potentially complex cognitive mechanisms, or information states that all individuals must have access to in order to have an adaptive excavation output.

The authors have addressed questions I had in the review process and the manuscript is now clear in its communication and conclusions.

The modelling approach is compelling, also allowing extrapolation to other group sizes and even other species. This to me is the main strength of the paper, as the answer to the question of whether it is younger or older ants that primarily excavate nests could have been answered by an individual tracking approach (albeit there are practical limitations to this, especially in the observation nest setup, as the authors point out). The analysis of the tunnel structure is also an important piece of the puzzle, and I really like the overall study.

We sincerely thank Reviewer #2 for the time and effort dedicated to our manuscript's detailed review and assessment.

Reviewer #1 (Recommendations for the authors):

Thank you for the modifications. I found much of the additional information very helpful. I do still have a few comments, which I will include below.

We thank the reviewer for this comment

The authors provide some additional citations for the model, however, the ODE in refs 24 and 30 is different from what the authors present here, and different from what is presented in ref 29. Specifically, the additional "volume" term that multiplies the entire equation. Can the authors provide some additional context for their model in comparison to these models as well as how their model relates to other work?

We thank the reviewer for this question. The primary difference between the logistic model (reference number: 24,30), and the saturation model (reference number: 29) is rooted in their assumptions on the scaling of the active number of ants that participate in the nest excavation and the nest volume.

The logistic growth model ( 𝑑𝑉/𝑑𝑡 = α𝑉(1-V/Vs) describes the excavation in fixed-sized colonies (50, 100, 200) through a balance of two key processes : (1) positive feedback (α𝑉), where the digging efficiency increases with the nest size, and (2) negative feedback (1-V/Vs), where growth slows as the nest approaches a saturation (Vs). The model assumes that the number of actively excavating ants is linearly proportional to the nest volume (V). This represents a scenario where a large nest contains or can support more workers, which in turn increases the digging rates. While this does not require explicit communication between individuals, ants indirectly sense the global nest volume through stigmergic cues, such as pheromone depositions, encounter rates, while ignoring individual differences in age.

In contrast, the saturation model (𝑑𝑉/𝑑𝑡 = α𝑉(1-V/Vs) assumes a constant number of ants is working throughout the excavation. The digging rate is therefore independent of the nest volume, this model imposes a different cognitive requirement ants must somehow assess the global nest slowing only due to the saturation term (1-V/Vs) as the nest approaches its target size. However, volume (V) and the overall number of ants in the nest. Thus, rather than relying on local cues, ants need more explicit communication or a sophisticated global perception mechanism that allows ants to sense the nest volume and the nest population to adjust the digging rates accordingly. Therefore, this model requires a more complex and less biologically plausible mechanism than the logistic model.

In our age-dependent digging model in the manuscript, we explicitly sum the contribution of each ant towards the nest area expansion based on its age-dependent digging threshold (quantified from fixed demographics experiments) the sum over Thus, the term ‘V’ in the ‘ 𝑉(1-V/Vs) takes the same effect as sum over all ants in the equation (2) of our manuscript; they describe how the total excavation rate scales with the number of individuals. Under the simplifying assumption that the number of ants is proportional to the nest volume ‘V’, and that all ants dig at a constant rate, our equation (2) in the manuscript reduces to the logistic equation ‘𝑉(1-V/Vs)’ This implies that each ant individually assesses the nest volume and then digs at a rate ‘(1-V/Vs)’.

Thus, we adopted the simpler model from the previously published ones, in which ants individually react to the local density cues and regulate their digging. This approach does not require a global assessment of the nest volume or the number of ants; a local perception of density triggers each ant’s decision to dig, likely modulated by the frequency of social contacts or chemical concentration, which serves as an indicator of the global nest area. The ant compares this locally perceived density to an innate, age-specific threshold. If the perceived local density exceeds its threshold (indicating insufficient area), it digs; otherwise, there is no digging. Thus, excavation dynamics in maturing colonies emerge from this collective response to local density cues, without any individual need to directly assess the global nest volume (V) or having explicit knowledge of the colony size (N).

As suggested by the reviewer, we have added these points to the discussion, contrasting the previously published models with our age-dependent excavation models (line numbers: 283-290) “In our study, we adopted the simpler version of previously published age-independent excavation models, where individuals respond to local stigmergic cues such as encounter rates or pheromone concentrations, which serve as a proxy for the global nest volume (24,30). We minimally modified this model to include age-dependent density targets. According to our age-dependent digging model, each ant compares this perceived local density to its own innate age-specific digging threshold as quantified from the fixed demographics experiments. If the perceived local density exceeds its age-dependent area threshold (indicating insufficient area), it digs; otherwise, there is no digging. This mechanism eliminates the need for cognitively demanding global assessment of the total nest volume or the overall colony population, a requirement for the saturation model (29)”.

I still find it a little concerning that the age-independent model, though it cannot be correct, fits the data better than the age-dependent modification. It seems to me the models presented in refs 24, 29, and 30, which served as inspiration for the one presented here, do not have any deep theoretical origin, but were chosen for "being consistent with" the observed overall excavated volumes. Is this correct, and if so, how much can/should be gleaned about behavior from these models? Please provide some discussion of what is reasonable to expect from such a model as well as what the limitations might be.

We thank the reviewer for the comment.

In our study, we make an important assumption, as described in the lines (line number : 161 - 164) of the manuscript, that ants rely on local cues during nest excavation, and individuals cannot distinguish between the fixed demographics and colony maturation conditions. This implies that the age-dependent target area identified in the fixed demographics experiments should also account for the excavation dynamics seen in the colony maturation experiments.

From the fixed demographics young and old experiments, we directly quantified that the younger ants excavate a significantly larger area than the older ants for the same group size. This age-dependent digging propensity is an experimental result, and not a model output.

We agree that the age-independent model fits the colony maturation experiments well, even though it's not a statistically better fit than the age-dependent model. However, the age-independent models in the references (24,29,30) fail to explain the empirically obtained excavation dynamics in the fixed demographics, young and old colonies. If indeed these models were true, then we would have observed similar excavated areas between the colony maturation, fixed demographics, young, and older colonies of the same size. Thus, the inconsistency of these models confirms that age-independent assumptions are biologically inadequate. These details are explicitly mentioned in lines (304 - 309).

We believe that our model’s value is in providing a plausible explanation for the observed excavation dynamics in the colony maturation experiments, and generating testable predictions (Figure 4. C, and 4.D, described in lines 199 - 216) about the percentage contribution of different age cohorts and queens to the excavated area from the colony maturation experiments. This prediction would not be possible with an age-independent model.

Minor comments:

Figure 2A: Please use a color other than white for the model... this curve is still very hard to see

We thank the reviewer for the comment. The colour is changed to yellow.

Figure 4A: Should quoted confidence intervals for slope and intercept be swapped?

Yes, we thank the reviewer for pointing this out. The labels for the slope and intercept were swapped. We corrected this in the current revised version 2.

Figure 5 D-F: Can the authors show data points and confidence intervals instead of bar graphs? The error bars dipping below zero do not clearly represent the data.

We thank the reviewer for the comment. We now show the individual data points from each treatment with the 95% Confidence Interval of the mean.

Changes from the earlier version to the revised version

Removed in the revised version: Abstract

Abstract

The ant nest serves as the skeleton of the ant superorganism. Similar to a skeleton, the nest expands as the colony grows and requires repair after catastrophic events. We experimentally compared nest excavation in colonies seeded from a single mated queen and allowed to grow for six months to excavation triggered by a catastrophic event in colonies with fixed demographics, where the age of each worker, including the queen, is known. The areas excavated by equal group sizes differed significantly between these conditions: heterogeneous populations in naturally growing colonies as well as cohorts of young ants dig larger areas than old ant cohorts. Moreover, we find that younger ants tend to dig slanted tunnels while older ants dig straight down. This is a novel form of age polyethism, where an ant’s age dictates not only her likelihood to engage in a task but also the way she performs the task. We further present a quantitative model that predicts that under normal growth, digging is predominantly performed by the younger ants while after a catastrophe all ants dig to restore lost nest volume. The fact that the nests of naturally growing colonies exhibit slanted tunnels strengthens this prediction. Finally, our results indicate how a colony’s demographic and physical history are sketched into the current structure of its nest.

Introduction

Nests occur across a range of biological scales and serve various functions such as providing refuge, storage, transport, mating, and micro-climate regulation (1–10). Among the diversity of nests those constructed by social insects are highly complex in terms of function, architecture, and construction mechanisms. These nests are created through digging or modifying a substrate and can be arboreal or underground (11–18). Previous studies have demonstrated that these structures emerge from the interaction between individual workers and the environment in a decentralized manner, showcasing an exemplary instance of biological self-organization (19–28).

Previous studies on nest construction and dynamics often addressed mature colonies of fixed size and over short time scales (24, 29–31). They show that for ant groups of fixed size nest volume grows logistically to a saturation volume that scales linearly with population size (24, 29, 30, 32–34). This led to the hypothesis that workers can perceive the density of ants within the nest and adjust their digging activity accordingly. This mechanism may hold in relatively small populations and simple architectural designs. However, for more complex structures where ants congregate in specific chambers, workers are less likely to assess the overall nest density (24, 29, 30). In such cases, nest excavation is enhanced by recruiting additional workers from the other areas of the nest or facilitated by humidity or pheromones (22, 24, 27, 29, 35–38).

Ant colonies are composed of ants that vary in age. It is well known that the task allocation of individual workers can change throughout their lifetimes (age polyethism) (39–46). However, it is unknown which of the work groups in a developing colony participate in the digging effort. Additionally, ant colonies are comprised of the reproductive caste and one or more worker castes that can divide labor among themselves (23, 39, 44, 45, 47, 48). Thus, colony demographics play a significant role in the development of nest structure, and its role needs to be considered.

One way to view the social insect nest is from a colony-superorganism perspective (2, 19, 20, 22, 26). Within this analogy, the social insect nest can be likened to the superorganism’s skeleton. Inspired by the processes of skeletal growth and bone healing after fracture we set out to find how the volume of an ant nest is regulated during ordinary colony development and following a catastrophic collapse. Specifically, we were interested in the following questions:

How does a colony regulate the co-expansion of the population and the nest as the colony develops from a single queen to tens of workers?

How does a colony recover nest volume following a catastrophic collapse?

Which ants participate in nest excavation and does this answer differ between the cases (a) and (b)?

We performed two sets of experiments to study how a Camponotus fellah ant colony regulates the volume of its nest. Camponotus fellah ants are native to the Near East and North Africa, particularly found in countries like Israel, Egypt, and surrounding arid and semi-arid regions, where they prefer to nest in moist, decaying wood, including tree trunks, branches, or stumps (49, 50). These ants live in monogynous colonies, ranging in size from a few dozen to several thousand individuals. Their nests are commonly built in a sand-loamy mix—a combination of sand, soil, clay, or gravel—providing structural stability and moisture retention (51). They are typically found under rocks, in the crevices of dried vegetation, or within dry, sandy soils, sometimes in areas with loose gravel (Materials and Methods).

The "Colony-Maturation" experiment observed the development of colonies for up to six months, starting from a single fertile queen and progressing to colonies with established worker populations. We used these experiments to track the coexpansion of the colony and its nest for an extended period (up to over six months). In another set of experiments, the "Fixed-Demographics" experiments, we monitored excavation by an ant group of determined size and age, which included the ant queen. We used these experiments to study the nest volume recovery following catastrophic events and the role of different age groups in the digging process.

Results

Nest design and experimental setup

All experiments were conducted within two-dimensional nests (Fig. 1, Materials and methods), nest sizes are reported as areas but these can be readily translated into volumes by multiplying by the thickness of the setup (0.8 cm). We performed two kinds of excavation experiments : colony maturation experiments that follow nest excavation alongside colony growth and fixed demographics experiments that track digging by already mature colonies.

Experimental system and nest architecture.

(A) We used large nest frames 80 ∗ 60 ∗ 0.8 cm3 to allow for long-term excavation. Nest frames are filled with homogenous fine sand and wet regularly for structural cohesiveness. Water infiltrates through the water column and keeps the sand moist and cohesive. (B and C) A typical nest at the start and termination of an experiment.

We initiated colony maturation experiments by introducing a single mated queen and several brood items (n = 5, across all experiments) at random positions on the soil layer of the nest. This allowed us to capture nest excavation and colony development from their very first stages. In all instances where a collapse resulted in the queen’s death or her being irreversibly trapped in the nest, the experiment was excluded from analysis starting from the point of the collapse, as such events did not reflect normal colony dynamics. Overall, colony maturation experiments (n = 22) spanned up to 6 months (see Table 1; Experimental details, Materials and Methods) during which the population size and nest area were recorded. The number of experiments decreased with time: for example, on day 130, we only had 13 ongoing excavation experiments (n = 9, collapses).

Fixed demographics experiments (n = 37) included colonies with known age demographics. In preparation for these experiments, we tracked the age of individual workers over the course of 10 months (Materials and Methods). We then extracted the queen alongside several workers of a given age and transferred them into the digging setup. Fixed demographics experiments spanned six conditions which varied by worker age (young vs. old workers) and group size (5, 10, &15 ants, including the queen). These experiments allowed us to track nest excavation as the ants were introduced into the setup and to further quantify the response to an induced catastrophic collapse.

Nest and colony development from the founding queen

Colony maturation experiments allowed us to track the concurrent growth of the colony and its nest.

Earlier version: In all colonies, population size increased logistically (Fig. 2 A) and saturated at a value that was highly variable across colonies (typically between 5 and 20 ants, see SI. Fig 1 and Table 1). As the population expanded so did the nest (Fig. 2 B). The nest excavation was commenced by the founding queen which dug a mean area of 23.8(± 1.7) cm2 (n = 22; unless noted otherwise, means are presented as mean ± standard error). Once the first workers emerged from the pupae the queen ceased digging (personal observation). However, these first workers did not dig (see kink in Fig. 2 B-C) possibly because the initial area dug by the queen led to low ant densities (number of ants/area) at this stage. As more workers emerged the area of the nest increased and varied considerably across nests (SI. Fig 1). For instance, halfway through the experiment duration, around day 95, the excavated area ranged from 38.41 cm2 to 173.28 cm2 with a mean excavated area of 100.51(± 11.54) cm2 (n = 15). At day 190, excavated areas ranged from 149.87 cm2 to 234.65 cm 2 with a mean excavated area of 197.47(± 13.8) cm2 (n = 6). The number of active experiments decreased over time since experiments were terminated as specified above.

Revised version: In all colonies, population size increased logistically (Fig. 2 A) and saturated at a value that was highly variable across colonies (typically between 5 and 20 ants, see SI. Fig 1 and Table 1). As the population expanded so did the nest (Fig. 2 B). The nest excavation was commenced by the founding queen which dug a mean area of 23.8(±1.7) cm2 (n = 22; unless noted otherwise, means are presented as mean ± standard error). Once the first workers emerged from the pupae the queen ceased digging (personal observation). However, these first workers did not dig (see kink in Fig. 2 B-C) possibly because the initial area dug by the queen led to low ant densities (number of ants/area) at this stage. As more workers emerged the area of the nest increased and varied considerably across nests (SI. Fig 1).

Excavation and population dynamics during colony maturation.

Earlier version: (A) & (B) Temporal change of the number of ants (A), and excavated area (B) starting from a single queen. Solid violet, and grey circles show the mean number of ants and excavated area across the 22 colony maturation experiment. Error bars depicted in black represent the standard error (here, and in all other figures). In (A), we fit a population growth model (white continuous line) using the number of ants as averaged over all colonies (Materials, and methods). (C) The average excavated area is plotted against the corresponding population size. The solid red line depicts a linear fit y = 10.71x, R2 = 0.97. The color bar represents the time in days. (D) The normalised cross-correlation coefficient between the digging rate, and population growth rate is shown. The digging rates are correlated with earlier (negative lags) population change with maxima at lag ‘− 10’, rather than with the current (zero lag) or future (positive lags). Solid squares correspond to the mean cross-correlation coefficient. The vertical green line highlights the lag time with the maximum mean cross-correlation coefficient.

Revised version: (A) & (B) Temporal change of the number of ants (A), and excavated area (B) starting from a single queen. Solid violet, and grey circles show the mean number of ants and excavated area across the 22 colony maturation experiment. Error bars depicted in black represent the standard error (here, and in all other figures). In (A), we fit a population growth model (yellow continuous line) using the number of ants as averaged over all colonies (Materials, and methods). (C) The average excavated area is plotted against the corresponding population size. The solid red line depicts a linear fit y = 10.71x, R2 = 0.97. The color bar represents the time in days. (D) The normalised cross-correlation coefficient between the digging rate, and population growth rate is shown. The digging rates are correlated with earlier (negative lags) population change with maxima at lag ’−10’, rather than with the current (zero lag) or future (positive lags). Solid squares correspond to the mean cross-correlation coefficient. The vertical green line highlights the lag time with the maximum mean cross-correlation coefficient.

Added in the revised version: For instance, halfway through the experiment duration, around day 95, the excavated area ranged from 38.41 cm2 to 173.28 cm2 with a mean excavated area of 100.51(±11.54) cm2 (n = 15). At day 190, excavated areas ranged from 149.87 cm2 to 234.65 cm 2 with a mean excavated area of 197.47(±13.8) cm2 (n = 6). The number of active experiments decreased over time since experiments were terminated as specified above.

Previous works on already mature colonies revealed a linear relation between population and nest size (24, 29). We find that this relationship persists in cases where colonies and nests are allowed to co-expand (Fig. 2 C). A linear fit (R-square - 0.97), indicates that the colonies regulate an area of approximately 11cm2 per each ant. We further quantified the relationship between the area per ant and experimental duration from the colony maturation experiments using two alternative definitions of day 1: (a) the start of the experiment and (b) the emergence of the first worker. We found that the area stabilized at 11.1 (±1) for scenario (a) and 11.6 (±1.15) for scenario (b) (SI Fig. 2 A, B). Deviations from this fit at early times result from the initial area excavated by the queen.

All the results stated so far deal with averages over different colonies. These do not allow us to address the causality of the regulation depicted in Fig 2 C. To test whether colony expansion triggered changes in nest-area or if it was the other way around, we tracked these concurrent processes one colony at a time. In all experiments, the initial population growth is followed, 15(±12) days later, by an increase in excavated area. Moreover, in some of the colonies these dynamics repeat themselves leading to a step-like pattern (SI Fig. 3). This staggered pattern suggests that the nest is excavated in separate digging episodes, possibly instigated by a population increase. Cross-correlation analysis between the digging rate and the population growth rate (Fig. 2 D) indeed reveals that an increase in the population size is followed by an increase in nest area on a time scale of one week.

Added in the revised version: Excavation dynamics for similar-age groups.

Added in the revised version: (A & B) The area excavated by the fixed demographics young and old colonies with a colony size of 5,10, and 15 ants. Solid stars and circles indicate the mean excavated area on the day of collapse. Diamonds and triangles indicate the mean area recovered after the manual collapse. (C) The mean area excavated by colonies of different group sizes from the colony maturation and fixed demographics experiments are shown. Squares (colony maturation), stars (fixed demographics young), and circles (fixed demographics old) correspond to the mean area excavated by colonies of a particular group size. Shaded error bars and error bars correspond to the SEM.

Our results suggest that population size controls nest area and not vice-versa. As newly-born ants join the colony the size of the nest increases accordingly. It is, however, not clear which ants do the actual sensing or digging. One possible explanation is that, although all ants are capable of digging, it is primarily the newly emerged ants who perform this task. In this case, nest expansion would lag behind colony growth due to two delays: first, the time needed for young ants to mature enough to begin digging, and second, the physical time required to excavate additional space (e.g., around 10 days). This mechanism could eliminate the need for ants to assess overall colony density, as each new group of active workers simply enlarges the nest as they become ready. An alternative possibility is that all ants, regardless of age, respond to increased density by initiating excavation. In that scenario, nest expansion would follow more immediately after the emergence of new individuals, making delays less prominent (24, 29, 30).

Fixed demographics experiments reveal age-dependent nest excavation

In "Colony-maturation" experiments ants emerge from their pupae inside the thin confines of the setup. This prevents us from marking them by age. To test how an ant’s age affects her propensity to dig, we conducted shorter "Fixed-Demographics" excavation experiments with fixed population sizes (5,10, or 15 ants including a queen) where all workers are of similar age (either "young” (40 ± 16 days) or "old” 171 ± 20) (SI Fig. 4). These ages represent functionally different life stages - the younger group had completed about 25% of their expected lifespan at the start of the experiment, while the older group had lived through most of theirs (50, 52). This 4-fold age difference allowed us to compare excavation behaviors across fundamentally different phases of adult life. These experiments (Materials, and Methods, Table 2) further allowed us to study excavation under conditions of extreme area shortage. Under our protocol, such conditions occur twice: as the ants are initially introduced to the setup and then again after we induce a collapse.

Added in the revised version: Age dependent nest excavation.

Added in the revised version: (A) The target area per ant is represented as a function of the age of individual ants (dotted line). Target areas decrease linearly with the ant age (y = −0.032x + 11.22, 95 % CI (Slope : (-0.035,-0.027), Intercept : (10.53,11.91)), R2= 0.96). The fixed target area (11.6cm2ant−1) calculated from the colony maturation experiment is shown with a solid blue line. (B) The mean excavated area as predicted by the age-dependent digging model is shown alongside the excavated area from the colony maturation experiments. Error bars in (A,& B) represent the SEM. (C) Total excavated area during normal colony growth as broken down by young ants, old ants, and the queen as predicted by the age-dependent model. The age threshold for young ants was set to be 56 days. Solid circles represent the area excavated by ants across 22 experiments. Boxes and median lines represent inter-quartile range and median values, and whiskers represent minimum and maximum values of data within 1.5-fold of the inter-quartile range. (D) The area excavated by each worker ant during normal colony growth as estimated by the age-dependent model.

Following the introduction of the ants into the setup, the nest area grew logistically and stabilized after about three weeks independently of age or group size (Fig. 3 A, B). We found a significant difference in the maximum area excavated by same-sized cohorts as a function of worker age, with the younger ants excavating larger areas when compared to the older ants (Fig. 3 A, B-Unpaired Wilcoxon signed-rank test significant for group size: 10; p = 0.014, 15; p = 0.042). Once the fixed-demographics nests reached a steady-state area, we artificially collapsed 25-30% of the steady-state area (star - Fig 3A, & circle - Fig 3B, Materials and Methods - Nest artificial collapse). We found that irrespective of age, ants dug after the collapse. Moreover, colonies ceased digging when they recovered (93 ± 3)% of the area lost by the manual collapse (diamonds, and triangles - Fig 3A, & 3B).

We compared the saturated excavation areas (pre-collapse) from fixed-demographics experiments (young and old groups) with those from colony maturation experiments of the same colony sizes (Fig. 3 C). We find that, for a given age cohort (young or old), the saturation areas increase linearly with colony size (GLMM, F(35,37); p < 0.0001) (Fig. 3 C, SI. Fig 7 A). The observed proportional scaling between excavated area and group size aligns with previous studies, even though those studies did not explicitly account for age demographics (24, 29, 30). After normalizing the pre-collapse excavated area by group size for both young and old colonies, we found no significant difference in area per ant across group sizes (SI. Fig 5. A). This indicates that the excavated area per ant remains relatively constant within each demographic group. Furthermore, the area excavated by the young cohorts was comparable to that excavated by naturally maturing colonies when they reached the same population size (Tukey’s HSD; group size: 5, p = 0.61; group size: 10, p = 0.46; group size: 15, p = 0.20). This holds even after excluding the area excavated by the queen from the colony maturation experiments (SI. Fig 5. B). Digging rates (change in excavated area per day) for all group sizes in fixed-demographics experiments (young and old ants) peak within the first 7 days, then asymptotically decay to near-zero levels as nest excavation approaches saturation (SI. Fig. 8).

Added in the revised version: Structural classification of nests.

Added in the revised version: (A, B, & C). Panels represent the nests excavated from the colony maturation (A), fixed demographics young (B), and old (C) nests in terms of the tunnel, wide tunnel, and chambers for a colony size 15. (D, E, & F) The contribution of tunnels and chambers to the total excavated area for a fixed group size across all the experimental conditions. Solid circles represent individual data points, and error bars represent 95% CI of the mean. (G, H,&I) Angular orientation of tunnels and chambers across experimental treatments.

Normalizing saturation areas and digging rates by the number of ants allowed us to obtain the single ant digging rate as well as the target area per ant (this is the reciprocal of ant density which, for short, we refer to as the target area from here on). We find that while the target area of an ant depends on her age (Fig. 4 A), the post-collapse maximum digging rates revert to the pre-collapse levels, revealing an age-independent maximum digging rate (SI Fig. 6 A, B), over the first 200 days of an ant’s life (SI. Fig 7 B).

Age-dependent model predicts minimal digging by older ants

In the previous sections, we showed that in fixed-demographics experiments, younger ants excavated a significantly larger nest area compared to older ants (Fig. 3 C). This difference emerged despite similar temporal patterns in digging rates across age groups, with excavation activity peaking within the first 7 days before asymptotically decaying as nest expansion approached saturation (SI. Fig. 8). We make a key important assumption in our model: ants rely on local cues during nest excavation, and individuals cannot distinguish between the fixed demographics and colony maturation conditions. This implies that the age-dependent target area identified in the fixed demographics experiments should also account for the excavation dynamics seen in the colony maturation experiments. This assumption facilitates a re-analysis of the colony maturation experiment, allowing us to estimate which ants are involved in nest excavation during normal nest growth and investigate its implications.

Removed in the revised version: Excavation dynamics for similar-age groups.

Removed in the revised version: (A & B) The area excavated by the fixed demographics young and old colonies with a colony size of 5,10, and 15 ants. Solid stars and circles indicate the mean excavated area on the day of collapse. Diamonds and triangles indicate the mean area recovered after the manual collapse. (C) The mean area excavated by colonies of different group sizes from the colony maturation and fixed demographics experiments are shown. Squares (colony maturation), stars (fixed demographics young), and circles (fixed demographics old) correspond to the mean area excavated by colonies of a particular group size. Shaded error bars and error bars correspond to the SEM.

Earlier version: To do this, we present a modified version of the constant density regulation model of nest excavation (24, 29, 30). In the constant density regulation model, all ants dig towards the same target area, regardless of their age. In contrast, based on our findings from the fixed demographics experiment, our modified model incorporated a dynamic component; the digging rate of each ant now depends on how similar the current area per ant (a(t), which is one over the density) is to her age-dependent target (Equation Eq. 1). Where aage is the target area of a specific age (Fig. 4 A) and the basal digging rate constant is(as above). The instances where the current area per ant (‘a’) is higher than the steady-state area per ant a(age) resulted in a negative digging rate. While there were very few instances where ants filled up the excavated regions they were extremely rare. Therefore, in the model, negative digging rates were taken to be zero. In the model, we assume that the ants rely on local cues to perform nest excavation, and they are unaware of the experimental condition (for example; fixed demographics, or colony maturation).

Revised version: To do this, we present a modified version of the constant density regulation model of nest excavation (24, 30). In the constant density regulation model, all ants dig towards the same target area, regardless of their age. In contrast, based on our findings from the fixed demographics experiment, our modified model incorporated a dynamic component; the digging rate of each ant now depends on how similar the current area per ant (a(t), which is one over the density) is to her age-dependent target (Equation Eq . 1). Where aage is the target area of a specific age (Fig. 4 A) and the basal digging rate constant is (as above). The instances where the current area per ant (‘a’) is higher than the steady-state area per ant a(age) resulted in a negative digging rate. While there were very few instances where ants filled up the excavated regions they were extremely rare. Therefore, in the model, negative digging rates were taken to be zero. In the model, we assume that the ants rely on local cues to perform nest excavation, and they are unaware of the experimental condition (for example; fixed demographics, or colony maturation).

Removed in the revised version: Age dependent nest excavation.

Removed in the revised version: (A) The target area per ant is represented as a function of the age of individual ants (dotted line). Target areas decrease linearly with the ant age (y = −0.032x + 11.22, 95 % CI (Intercept : (−0.035,-0.027), Slope : (10.53,11.91)), R2 = 0.96). The fixed target area (11.6cm2ant−1) calculated from the colony maturation experiment is shown with a solid blue line. (B) The mean excavated area as predicted by the age-dependent digging model is shown alongside the excavated area from the colony maturation experiments. Error bars in (A,& B) represent the SEM. (C) Total excavated area during normal colony growth as broken down by young ants, old ants, and the queen as predicted by the age-dependent model. The age threshold for young ants was set to be 56 days. Solid circles represent the area excavated by ants across 22 experiments. Boxes and median lines represent inter-quartile range and median values, and whiskers represent minimum and maximum values of data within 1.5-fold of the inter-quartile range. (D) The area excavated by each worker ant during normal colony growth as estimated by the age-dependent model.

The total excavation rate for a colony that includes ′N′ ants of different ages can be obtained by summing over the area contributions (a) from each ant: Since we recorded the population sizes in all colony maturation experiments, and mortality events were rare - we could readily deduce the entire age demographics of the colony for each experiment (SI section: Age quantification from colony maturation experiments, Table 3). Using this data, we could integrate Equation 2, over each experiment to obtain the total excavated area.

The models’ results and their comparison to colony maturation experiments are shown in Fig. 4 B. In both model and colony maturation experiments, ’day 1’ corresponds to the day on which the first workers emerge. The area at this day corresponds to the area excavated by the queen. We find that the age-dependent model fits the excavation dynamics observed in the colony maturation experiments. Furthermore, we simulated an age-independent density model (SI Fig. 9) using a constant target area per ant (11.6cm2ant−1, SI Fig. 2), as obtained from the colony maturation experiments, rather than an age-dependent area target that decreases with age (Fig. 4A). While the age-independent model provides a reasonable fit for the colony maturation experiments, it fails to align with the results of the fixed demographics experiments. This discrepancy arises because the fixed demographics experiments with young and old groups revealed a significant difference in the excavated area, despite having the same group size. This finding demonstrates that the area per ant is not a fixed value, as assumed in the age-independent density model. Notably, the target area fitted for the age-independent model closely approximates the empirically measured age-dependent target when extrapolated to very young ants (Fig. 4A). This raises the hypothesis that, in the colony maturation experiments, the youngest ants are responsible for the majority of the digging. Thus, in light of our assumption that ants rely on local cues and do not distinguish between experimental contexts, these findings direct us towards an age-dependent model to explain the excavation dynamics in the colony maturation experiments.

Since the target area is age-dependent and ants in the colony vary in age, not all ants perform the same amount of digging. By integrating equation 2 from the age-dependent model, we estimated the contribution of each age group to the total excavated area. We set an age threshold at 56 days (mean + standard deviation of the age of the younger cohort) in the age-dependent model to categorize the excavated area contribution between the young and old ants. This allowed us to quantify the area dug by the young ants, old ants, and the queen in terms of their contribution to the overall excavated area. We hypothesize that, as young ants have a lower digging threshold, they would perform the majority of the digging. Indeed, our model suggests that after the initial excavation by the queen, younger ants performed the vast majority of the digging (Fig. 4 C). The area excavated by older ants (age of 57 days and above) was nearly negligible. Thus, the age-dependent model provides indirect evidence indicating that younger members of the colony predominantly carry out nest excavation.

Earlier version: Furthermore, using the age-dependent model, we quantified the area dug by each ant for the colony maturation experiment (Fig. 4 D), and analyzed the percentage contribution of each histogram bin to the excavated area (provided in Table: 4). We found that 73.76 % of the total excavated area ranged between 5 cm2 and 13 cm2. This indicates that most ants dug within this range rather than exhibiting extreme variations. Additionally, the mean excavation amount is 7.84 cm2, with a standard deviation of 3.44 cm2, meaning that most values fall between 4.4 cm2 and 11.28 cm2, which aligns well with the 5–13 cm2 range. Since the majority of the excavation is concentrated within this narrow interval, and the mean is well centered within it, this suggests that ants excavated more or less the same amount, rather than forming distinct groups with highly different excavation behaviors.

Revised version: Furthermore, using the age-dependent model, we quantified the area dug by each ant for the colony maturation experiment (Fig. 4 D), and analyzed the percentage contribution of each histogram bin to the excavated area (provided in Table: 4). We found that 73.76 % of the total excavated area ranged between 5 cm² and 13 cm². This indicates that most ants dug within this range rather than exhibiting extreme variations. Additionally, the mean excavation amount is 7.84 cm², with a standard deviation of 3.44 cm², meaning that most values fall between 4.4 cm² and 11.28 cm², which aligns well with the 5–13 cm² range. Since the majority of the excavation is concentrated within this narrow interval, and the mean is well centered within it, this suggests that ants excavated more or less the same amount, rather than forming distinct groups with highly different excavation behaviors.

These findings suggest that while all ants can dig and will do so following catastrophic events (see Fig. 3), the regular nest expansion is primarily carried out by young ants.

Structural classification of nests

The colony maturation and fixed demographics experimental conditions differed not only in the total volume of the nest but also in its architecture. We compared nest architectures by segmenting raw nest images into chambers and tunnels (see SI Section: Nest Skeletonization, Segmentation, and Orientation). Chambers were identified as flat, horizontal structures, while tunnels were narrower and more vertical in orientation (see SI Fig. 10, SI Section: Nest Skeletonization, Segmentation, and Orientation) (13, 30, 53, 54)).

We found that the colony maturation colonies exhibited a “top-heavy” design (13, 55, 56): large chambers (marked in Yellow in Fig. 5 A) comprise most of the nest area at the top part of the nest (56). Tunnels generally branch out from the large chambers and mostly lead to the bottom parts of the nest. The wide tunnels are primarily seen originating from the tunnels and are oriented parallel to the chambers (Fig. 5 A). In the fixed demographics experiments, the nest structure mostly lacked large chambers, unlike the colony maturation experiments (Fig. 5 B, C, D, E, and F). This could be attributed to the absence of a young queen prone to dig and/or to the absence of brood. We analyzed the orientation of all tunnels and chambers relative to the direction of gravity. We found that in the colony maturation and fixed-demographics young experiments, the bulk of tunnels and wide tunnels are oriented between (30° - 60°). Surprisingly, in the fixed-demographics old colonies, the orientation of tunnels ranged from (70° - 90°) (Fig. 5 G, H; SI Fig. 11). Thus, the tunnels and wide tunnels in the colony maturation experiments are similar to the fixed demographics of young colonies but not the old ones.

Removed in the revised version: Structural classification of nests.

Removed in the revised version: (A, B, & C). Panels represent the nests excavated from the colony maturation (A), fixed demographics young (B), and old (C) nests in terms of the tunnel, wide tunnel, and chambers for a colony size 15. (D, E, & F) The contribution of tunnels and chambers to the total excavated area for a fixed group size across all the experimental conditions. Error bars represent the SEM. (G, H,&I) Angular orientation of tunnels and chambers across conditions.

Maturing colonies are composed of ants of different ages. Nevertheless, colony maturation experiments displayed tunnel orientation similar to those in the nest of young fixed demographics cohorts. This supports our prediction of the age-dependent model that younger ants primarily undertake digging tasks within maturing colonies.

Discussion

Group size is a key determinant in the behavior of ant colonies. Groups of different sizes differ in their foraging strategies (57–61), division of labor (2, 42, 45–48, 62, 63), and communication networks (64–67). It is less clear what happens in the expansion period that follows colony foundation, when group size is constantly on the rise. A growing colony may exhibit collective patterns that differ from a fixed group of the same size. This is especially relevant if the system preserves long-term memories such as the stigmergic alteration of the environment evident in nest construction (22, 30, 68–71).

To date, the overwhelming majority of studies on nest excavation in ants were short-term experiments performed on a mature group of fixed size (24–26, 29–31, 35, 56, 71). As we have shown, these short term studies likely address construction in emergency scenarios where there is a huge deficit in available space. Studying animal behavior in ethologically relevant conditions is most valuable for being able to correctly capture the nature of these animals and is crucial for deciphering the underlying meaning of their behavior. Furthermore, testing animals in the wrong context/conditions could lead to misleading conclusions (72–77).

The current study follows ant colonies over extended periods as they mature from a single reproductive individual. We find significant quantitative and qualitative differences between nests constructed during the colony maturation experiments, which mimic natural colony development from a single queen, and nests constructed in the context of an emergency. The differences include the nest volume, structural layout, tunnel orientation, and the demographic identity of ants involved in excavation. Colonies from the colony maturation experiments excavated the largest nest volumes for a given group size, followed by fixed-demographics young and then old colonies. Chambers were consistently present in colony maturation nests, sparse in young, and nearly absent in old colonies. The orientation of both tunnels and wide tunnels was similar in the colony maturation and fixed-demographics young experiments, but it was different in the old colonies.

A main advantage of social insects is their ability to allocate different tasks among individuals. A well-known example includes age-polyethism wherein ants perform age-specific tasks (2, 42, 45–48, 62, 63). Response thresholds that vary with age have been suggested as a mechanism for age polyethism (23, 78–82). Previous studies of nest excavation indeed rely on a threshold mechanism that prompts ants to dig extra space when the nest becomes too crowded (24, 29, 30, 35). Our results reveal that this target area of the group decreases linearly with age, such that young ants are more sensitive to space shortages. Thus, under normal growth conditions, nest volume remains proportional to the number of ants, ensuring a relatively constant area per ant. As a result, density thresholds for older ants are not exceeded, and excavation is predominantly carried out by younger ants. These findings have significant implications for the dynamics of nest volume.

In our colony maturation experiments, we found that area per ant was highest when the workers were youngest, with values around 11.1–11.6 (±1–1.15). This aligns with observations from naturally growing nests, where newly eclosed ants dominate the population and nest volumes are relatively large. Supporting this, fixed-demographics experiments showed that the area excavated per ant declines linearly with worker age, indicating that the youngest ants contribute most to excavation. Notably, the target area we fit for the age-independent model (11.6 ± 1.15) closely matches the extrapolated value for very young workers (Fig. 4A), reinforcing the idea that young ants are the primary excavators during early colony growth. In contrast, during events like collapses or displacement, when space is urgently needed, ants of all ages participate in excavation.

We highlight that Camponotus fellah queens can live up to 20 years, while worker lifespan is significantly shorter, though exact data for this species are lacking (45, 50, 52). However, studies on related Camponotus species (83) suggest that workers can live several months (≈ 1 year for workers, and ≈ 2 years for majors; personal observation). Since our observed worker lifespan (≈ 200 days) falls within this range, it likely represents a substantial portion of their natural lifespan, justifying their classification as "young", and "old".

Added in the revised version: In our study, we adopted the simpler version of previously published age-independent excavation models, where individuals respond to local stigmergic cues such as encounter rates or pheromone concentrations, which serve as a proxy for the global nest volume (24, 30). We minimally modified this model to include age-dependent density targets. According to our age-dependent digging model, each ant compares this perceived local density to its own innate age-specific digging threshold as quantified from the fixed demographics experiments. If the perceived local density exceeds its age-dependent area threshold (indicating insufficient area), it digs; otherwise, there is no digging. This mechanism eliminates the need for cognitively demanding global assessment of the total nest volume or the overall colony population, a requirement for the saturation model (29).

Earlier version: Our age-dependent model demonstrates that the digging behavior in Camponotus fellah is governed by a basal digging rate constant (r) modulated by the age-dependent feedback. Crucially, we show that after a collapse, the maximum digging rates return to their pre-collapse levels, suggesting that this basal rate ‘r’ represents an age-independent ceiling on how fast ants can dig, regardless of age or context (SI Fig. 6 A, B). Previous studies have demonstrated both homogeneous and heterogeneous workload distribution, with varying digging rates among ants (24, 29, 30, 35). Studies showing heterogeneous workload distribution relied on continuous individual tracking of ants to quantify digging rates (35). However, this approach was not feasible in our current design due to the experimental durations of both our colony maturation and fixed demographics experiments. Additionally, sample size requirements naturally limited our ability to conduct continuous individual tracking during nest construction in our study. Thus, based on empirical measurements from our fixed-demographics experiments and supported by the age-independent post-collapse digging rates, we adopted a constant basal digging rate for simulating our age-dependent model—an assumption aligned with both prior literature and the collective dynamics observed in our system (24, 29, 30).

Revised version: Our age-dependent model demonstrates that the digging behavior in Camponotus fellah is governed by a basal digging rate constant (r) modulated by the age-dependent feedback . Crucially, we show that after a collapse, the maximum digging rates return to their pre-collapse levels, suggesting that this basal rate ’r’ represents an age-independent ceiling on how fast ants can dig, regardless of age or context (SI Fig. 6 A, B). Previous studies have demonstrated both homogeneous and heterogeneous workload distribution, with varying digging rates among ants (24, 29, 30, 35). Studies showing heterogeneous workload distribution relied on continuous individual tracking of ants to quantify digging rates (35). However, this approach was not feasible in our current design due to the experimental durations of both our colony maturation and fixed demographics experiments. Additionally, sample size requirements naturally limited our ability to conduct continuous individual tracking during nest construction in our study. Thus, based on empirical measurements from our fixed-demographics experiments and supported by the age-independent post-collapse digging rates, we adopted a constant basal digging rate for simulating our age-dependent model—an assumption aligned with both prior literature and the collective dynamics observed in our system (24, 30).

Although we find that a constant density model provides a reasonable fit for the excavated area from the colony maturation experiments, the results from the fixed demographics experiments suggest otherwise. Despite having the same group size, young and old colonies exhibited a significant difference in the excavated area, which would not be expected if a constant density model were in effect. Thus, in light of our assumption that ants rely on local cues and do not distinguish between experimental contexts, these findings direct us toward an age-dependent model for excavation behavior.

Digging behaviors that vary across the ants’ life time influence not only the nest volume, but also its structure. During colony founding, excavation is performed solely by the queen, who typically creates several large chambers near the nest surface. These chambers, consistent with the density-target threshold model, delay the emergence of the first workers, who later assume the digging role. This transition shifts nest architecture from the initial top-heavy design to a more balanced structure as the excavation progresses. In particular, tunnel orientation varies with the age of the diggers, revealing a novel form of age polyethism. While previous studies focused on age-based task allocation (84–90), our results demonstrate that the nature of the task and its execution itself can shift with age. This behavioural plasticity introduces a new regulatory mechanism, distinct from the case-based performance differences driven by polymorphic species (27, 31).

Our results suggest that the nest reflects the cumulative effects of past events and conditions experienced by the colony. For example, in line with our experimental findings, if the colony expanded during a period with many young ants, more slanted tunnels are expected, indicating accommodation for a growing population. Conversely, nest construction initiated by mature ants or a catastrophe leading to nest area reduction would result in newly constructed tunnels aligning more with gravity, running straight down. Thus, the nest structure could be treated as a historical record of the colony.

To conclude, ant colony demographics play a pivotal role in nest excavation. They affect the identity of the ants that do the digging, the dynamics of nest growth, the volume of the nest per population size, and the chamber and tunnel architecture of the nest. We find that age-polyethism provides a major means by which the colony exerts control over its nest in different scenarios.

Earlier version: Materials and methods Study species

Revised version: Materials and methods

Added in the revised version: Study species

We chose to work with Camponotus fellah, a common species in Israel of the ground-dwelling carpenter ants (7.8–17.2 mm body length, Ionescu-Hirsch, 2009). These ants are omnivorous, monogynous, and singly mated, with colony sizes ranging from tens to thousands of workers (51, 52). We collected newly-mated queens on the Weizmann Institute of Science campus after nuptial flights in March 2018. For the colony maturation experiments, we used only queens with pupae, that had already nurtured larvae to increase survival rates. We used colonies of group sizes (5,10, and 15 ants) for the fixed demographics experiments. We housed the single queen and fixed colonies in plastered nest containers and environmental conditions (climate-controlled chamber under controlled humidity (65%), temperature (27◦ C), and a light-dark cycle of 12 hours) before the experiments. We supplied the colonies with a food mixture (tuna, honey, eggs, vitamin mix, and water), water, and a 20 % sucrose solution.

Nest structure

The artificial nest structure consists of a 2D Plexiglas frame 60 cm wide, 80 cm high, and 8 mm deep made from 5 mm thick cast acrylic sheets. Six support blocks (2 cm) held the front and back frames together. We equipped the nests with a 5 cm wide water column connected by a series of 0.4 mm holes in the bottom that allow water transfer but not sand grains (Fig. 1 A). The top part of the frame was coated with Fluon (Sorpol) to prevent the escape of ants. A removable plexiglass lid perforated with numerous small (1 mm) holes for ventilation covers the top of the nest.

Nest preparation

Each nest was filled 70 cm high with homogenous fine-grained sand (AccuSand 30 − 40, Unimin Co, USA) with grain sizes of 425-600 µm. Before the experiments began, water infiltrated gradually from the water column to the sand, slowly saturating the entire sand column from the bottom. Apart from wetting the sand, this action rearranged the sand grains in even more compact packaging and lowered the chance of spaces opening up in the column. On top of the sand layer, we placed a 1-2 cm layer of soil on the sand column, thus significantly reducing evaporation and the subsequent collapses. The soil was collected from the Weizmann campus orchard and sieved through coarse (2 mm) and fine (0.5 mm) sieves (Haver & Boecker, Germany), and was sterilized. We manually made a 30 mm long and 8 mm wide hole in the middle of the solid layer (seeding point) to induce a digging behavior.

Nest maintenance

We provided food to the ants ad libitum through three separate tubes containing water, 20% sucrose water, and protein food. The protein mixture included egg powder, tuna, prawns, honey, agar, and vitamins. Each of the three tubes was filled with 5 ml of their respective contents and sealed with a cotton stopper to prevent overflow. The tubes were positioned at a slight angle and connected using a custom-made plexiglass adapter to facilitate the flow of liquids. These tubes were replenished once depleted, and regularly replaced once the nest maintenance was carried out bi-weekly.

Experimental overview

The study included 22 colony maturation and 37 fixed demographics experiments (Fixed demographics young, group size 5, n = 4, group size 10, n = 6, group size 15, n = 6); Fixed demographics old (group size 5, n = 6, group size 10, n = 7, group size 15, n = 8) of ant nest excavation. We placed the founding queen on top of a 60 cm wide, 70 cm high, and 8 mm thick soil-sand column in each nest. We photographed the nest development every 1-3 days. During each photography session, four pictures of the nest were taken, with a 2-second interval between each. The experiments were terminated once the queen died or became trapped following a collapse.

Fixed demographics colonies

Earlier version: Queens used for the fixed demographics experiments were collected during the nuptial flight of May 2020. Forty-five newly mated queens were reared for the fixed demographics experiments and 37 colonies that attained the required colony sizes needed for the fixed demographics were chosen for the experiments. The colonies were maintained with a 12-hour light/dark cycle, room temperature of 25± 2oC, and 60% humidity. Each freshly born ant was marked using a unique three-color combination of acrylic paint on the head, thorax, and abdomen to determine the age. We checked each colony at least once in 5 days to check whether new workers emerged. The marked worker’s age was recorded and updated in due course. In the fixed demographics experiments, we separated colonies into two different demographics: Young and Old. The young and old colonies had workers with an average age of 40 ± 16 days and 171.56 ± 20 days, respectively. Error represents the standard deviation of the mean.

Revised version: Queens used for the fixed demographics experiments were collected during the nuptial flight of May 2020. Forty-five newly mated queens were reared for the fixed demographics experiments and 37 colonies that attained the required colony sizes needed for the fixed demographics were chosen for the experiments. The colonies were maintained with a 12-hour light/dark cycle, room temperature of 25 ± 2°C, and 60% humidity. Each freshly born ant was marked using a unique three-color combination of acrylic paint on the head, thorax, and abdomen to determine the age. We checked each colony at least once in 5 days to check whether new workers emerged. The marked worker’s age was recorded and updated in due course. In the fixed demographics experiments, we separated colonies into two different demographics: Young and Old. The young and old colonies had workers with an average age of 40 ± 16 days and 171.56 ± 20 days, respectively. Error represents the standard deviation of the mean.

Experimental conditions

Earlier version: We maintained a stable environmental condition in the experimental room by maintaining a 12-hour light/dark cycle and room temperature of 25 ±2oC. Daily variations were 2-3 degrees, usually 24 − 27oC. More significant deviations ranging at most 20− 30oC did occur sporadically due to technical problems, but they did not last more than a few hours at a time. Humidity was more challenging to maintain, and it fluctuated in the range of 40 70%, averaging around 60%.

Revised version: We maintained a stable environmental condition in the experimental room by maintaining a 12-hour light/dark cycle and room temperature of 25 ± 2°C. Daily variations were 2-3 degrees, usually 24 − 27°C. More significant deviations ranging at most 20 − 30°C did occur sporadically due to technical problems, but they did not last more than a few hours at a time. Humidity was more challenging to maintain, and it fluctuated in the range of 40 − 70%, averaging around 60%.

Nest photography - Data collection

A photography sequence of different exposures recorded the nest development in a measurement. The nests were placed one meter away from the camera against a diffuser-covered LED backlight. Images were initially taken with Canon EOS 800D DSLR cameras with a 24 MP resolution and fitted with EF-S 18-55mm f/4-5.6 IS STM lenses. We controlled the cameras remotely using a USB cable through the ‘DigiCamControl’ software (Istvan, 2019) with the photography sequence programmed using MATLAB CameraController (Serge, 2019). The photography sequence consisted of four photos taken at 2-second intervals.

Nest artificial collapse

To induce manual collapses, we started by quantifying the nest area from the start of the experiment, as described in the Image Processing and Excavated Area section to get the steady-state nest area. Once, we made sure that the excavated area approached a steady state, we fixed that 25–30 % of the excavated area should be collapsed. This calculation was based on image analysis, that provided the area to be collapsed in square centimeters. To ensure accuracy, we used graph paper to trace the entire nest area, allowing us to map the overall structure to a scale that facilitated precise collapse induction. We then compared the nest area obtained from image processing with the graph paper tracings to verify consistency in both the steady-state area and the area designated for collapse. Next, we fixed the graph paper onto the Plexiglas of the nest, against a light background, enabling precise alignment of the nest skeleton with all excavated structures. The collapsed area was marked on the graph paper, ensuring precise targeting for collapses. The collapse was induced in the deepest excavated region of the nest, in the direction of gravity, using a small cylinder with a 2 mm diameter. We specifically chose to collapse the deepest region to prevent further structural degradation that could lead to unintended collapses. Additionally, this approach ensured that ants were not irreversibly trapped within collapsed sections, allowing them to continue excavating the nest without obstruction.

Image processing and Excavated area

The images were post-processed using the MATLAB Image Processing Toolbox. The first step was image registration, where we aligned nest images with the first image taken before introducing ants. The algorithm registers the images through a projective geometric transformation using the six support blocks in the frame as control points. After registration and cropping, we binarized the images to show only the excavated areas of the sand column. The first binarization step involved subtracting the initial image from the processed image to reduce glare and noise. We then converted the RGB image into an intensity image by choosing the maximum value between the green or blue channels. Next, we used adaptive thresholding to produce BW (black and white) images of the nest architecture in which white pixels represented the excavated areas. Lastly, we made manual corrections to mark known excavated zones, which were not evident in the images, or remove other apparent noises. The excavated area refers to each Black and White image’s white pixel area.